---
title: "Digital Love"
output: 
  learnr::tutorial:
    progressive: true
tutorial:
  id: "Week_5_Extraction"
runtime: shiny_prerendered
description: "Introducing the basics of extraction in R"
---
```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
knitr::opts_chunk$set(echo = FALSE)
```

```{r info-section, child = system.file("child_documents/info_section.Rmd", package = "tutorial.helpers")}
```

```{r copy-code-chunk, child = system.file("child_documents/copy_button.Rmd", package = "tutorial.helpers")}
```

## Extraction Basics

A core part of data exploration and analysis is the ability to subset (or filter) your data into smaller units. Key to this is the idea of data *extraction*, which focuses on retrieving specific subsets of your data. These subsets may be a single or multiple values that meet a specific criteria of interest. You may additionally want to extract entire columns (variables) from a data frame. You can further use extraction to *add new data* to your data frame.

There are two key operators useful for extracting data from a data frame:

* `[]` extract specific values based on a reference to a row, column, or both.
* `$` extract or *add* entire columns (variables).

## Extracting data

Remember that data frames are just special data structures (lists) composed of different atomic vectors (columns) or the same length. The `[]` operator works on vectors and thus works on columns in a data frame. To start, run the code below to create a toy vector:

```{r demo_extraction, exercise = TRUE}

#create and print a vector
toy_vector <- seq(from = 10, to = 100, by = 10)
toy_vector

```

We can use the `[]` operator to *extract* values from vectors. Let's say you want to extract the third value in toy_vector. Run the code below.

```{r demo_extraction_2, exercise = TRUE,exercise.setup='demo_extraction'}

#extract third value from the vector
toy_vector[3]

```

Notice two things. First, the operator comes *after* the object it is operating on. Second, the location of the value we are interested in is specified via a numeric value inside the brackets. This extraction is not restricted to one value; we can actually extract multiple values from an object (here a vector) by *using a vector* to specify the sequence of values to extract. Run the code below

```{r demo_extraction_2_5, exercise = TRUE,exercise.setup='demo_extraction'}

#extract multiple values from the vector
toy_vector[c(3,5,7)]

```

Now let's say you are creating a data frame. When you do this, you are moving from a one-dimensional data structure (vector) to a two-dimensional data structure that has both rows and columns. The `[]` operator enables you to find specific rows *and* columns in a data frame. In this way, we can extract from more than just atomic vectors.

First, create a toy data frame using the code below. The two variables are the type of biome ("biome") and average carbon uptake (photosynthesis; "mean_C_uptake").

```{r demo_extraction_3, exercise = TRUE}

#define vectors
mean_C_uptake <-  c(700,400,200)
biome <- c("forest",'grassland','desert')

#create data frame from vectors
biome_data <- data.frame(mean_C_uptake,biome)

#print data frame
biome_data

```

Now let's say you want to *extract* the value in the third row in the second column of the biome_data data frame you created. You can use the `[]` operator to do this. Run the code below.

```{r demo_extraction_4, exercise = TRUE,exercise.setup='demo_extraction_3'}

#extract the value from the third row and second column
biome_data[3,2]

```

**The first value inside the brackets always refers to the row number, whereas the second value inside the brackets always refers to the column number**. By providing the row and column number, the `[]` operator can map to a specific value (cell) in the biome_data data frame.

*In the code chunk below, play around with what happens if you add [2,] or [2] or [,2] to the end of biome_data; what do you notice?*

```{r demo_extraction_5, exercise = TRUE,exercise.setup='demo_extraction_3'}

#play around with what value go in the brackets
biome_data[]

```

The `$` operator is particularly good if you want to extract or reference specific columns (variables). This operator gets "sandwiched" in between the data frame name and column name we are referencing. Run the example code below.

```{r demo_extraction_6, exercise = TRUE,exercise.setup='demo_extraction_3'}

#extract the biome column
biome_data$biome

```

Change the value of biome to mean_C_uptake in the code chunk above to see what happens.

Note that the `$` operator is less flexible than the `[]` operator and is specific to working with data frames. However, the `$` operator can be quite useful if you want to quickly calculate a summary statistic on a particular variable. Consider the `max()` function, which finds the maximum value across a numeric vector. If you wanted to quickly find the maximum carbon uptake in the entire biome_data data frame, it would like like this.

```{r demo_extraction_6_5, exercise = TRUE,exercise.setup='demo_extraction_3'}

#find the maximum carbon uptake
max(biome_data$mean_C_uptake)

```

Why does this work? It works because you are performing a calculation (using a function) on an atomic vector (a column), albeit one that has been *extracted* from a data frame. To reaffirm this, you can use the `is.atomic()` function on the extracted column to assess whether or not it is 'true' that what is being passed into the function is an atomic vector. The function returns a logical TRUE or FALSE. Run the code below.

```{r demo_extraction_6_75, exercise = TRUE,exercise.setup='demo_extraction_3'}

#determine whether the extracted column is an atomic vector
is.atomic(biome_data$mean_C_uptake)

```

Looks like it is "TRUE", which means that that column is indeed an atomic vector like we would expect for a tidy data frame.

## Adding new variables to a data frame

In addition to extracting a column (variable) from a data frame, you can also use the `$` operator to add new columns to a data frame.

Let's go back to the biome_data data frame you created in the previous section. Below is a refresher of what that looks like; run the code below.

```{r demo_extraction_7, exercise = TRUE,exercise.setup='demo_extraction_3'}

#view the data frame
head(x = biome_data)

```

Now let's say we want to add a column that will provide additional information pertaining to what type of plants are dominant within each biome. We can use the `$` operator to do so. Run the code below.
```{r demo_extraction_8, exercise = TRUE,exercise.setup = 'demo_extraction_3'}

#add a plant variable
biome_data$dominant_plants <- c("tree","grass","shrub")

#print the data frame
biome_data

```

More often you will want to calculate a new variable using existing variables in the data frame. You can do this in the same way you would do it for any numeric vector. Let's say you want to convert the mean_C_uptake column from grams (the units it is in) to kg. You can do this by 1) following the basic approach of the previous code chunk and 2) simply multiplying the mean_C_uptake column by 0.001 (which, in essence, is a vector of 1). This is just recycling like you have seen for vectors but now just within a column in a data frame. Run the code below

```{r demo_dataframe_9, exercise = TRUE,exercise.setup = "demo_extraction_8"}

#convert from grams to kg
biome_data$mean_C_uptake_kg <- biome_data$mean_C_uptake*0.001

#print data frame
biome_data
```

Calculating new variables from existing variable can also include equations that include using many variables in a data frame, such as calculating the difference between or ratio of two numeric variables.

## Filter data by a condition

Finally, we can bring the `[]` and `$` operators together to filter data. Both "filter" and "subset" are often used interchangeably and tend to imply a reduction in the number of rows in a data frame. Filtering data involves excluding data (rows) according to a set criteria or condition. In R, there are `comparison operators` that can help define relationships and establish filtering criteria.

* `==`	Denotes complete equality (exact match). For example,	x == y.
* `!=`	Denotes "is not equal to". For example,	x != y.
* `>` Denotes	"greater than". For example,	x > y.
* `<`	Denotes "less than". For example,	x < y.
* `>=` Denotes "Greater than or equal to". For example,	x >= y.
* `<=` Denotes "Less than or equal to". For example,	x <= y.

You can use these comparison operators with extraction operators to filter data by whatever criteria you set. For example, if you wanted to filter to just the row of data with the forest biome, it would look like this. Run the code below.

```{r demo_extraction_10, exercise=TRUE,exercise.setup='demo_extraction_8'}

#filter data frame to just the forest group
biome_data[biome_data$biome == 'forest', ]

```

Beginning from the inside out, the code first extracted the entire biome column using the `$` operator. It then designated an *exact* match via the `==` operator. All of this is occurring within the brackets of the `[]` operator, which if you remember has both a rows and column designation for extracting from data frames. Because the code is setting the filtering criteria only for rows, it will return the entire row that meets that criteria, along with all columns (since no criteria was set for columns). Think about why it would be most useful to return all rows (and thus a smaller version of the initial data frame) instead of just a single column or a single value. A plain language version of the code is: *filter the biome_data data frame down to rows where the biome variable is "forest".*

*Note that it might be the case that you will want to create a new, filtered data frame from another data frame, while retaining the original data frame. To do this, you can assign the new, filtered data frame a different name using the `<-` operator.*

In the same way, you can filter according to a numeric criteria. 
```{r demo_extraction_11, exercise=TRUE,exercise.setup='demo_extraction_8'}

#filter data frame to just the maximum 
biome_data[biome_data$mean_C_uptake > 200, ]

```
*All* rows of data that have mean_C_uptake greater than 200 are returned.

A final operator you can use is the `%in%` operator, which assess whether elements in one vector are present in another, and will return a logical TRUE or FALSE whether or not that is the case. You can use it to quickly filter data according to a set criteria. 
```{r demo_extraction_12, exercise=TRUE,exercise.setup='demo_extraction_8'}

#filter data frame to rows with grasslands and desert biomes
biome_data[biome_data$biome %in% c('grassland','desert'),]

```

Now on to some practice questions!

## Problem set

### Q1: Use the `[]` operator to extract the third, seventh, 10th, 13th, 256th, 1000th, 3000th, and 3500th values from the atomic vector below. The result should be a numeric vector containing all the values. First, run the code chunk below containing this vector.

```{r setup_problem_1, exercise = TRUE}

code <- seq(from = 300,to = 3000, by = 0.5)

```

Now use the `[]` operator to extract the third,seventh,10th,13th, 256th,1000th, 3000th, and 3500th values from "code".

```{r problem_1,exercise = TRUE,exercise.setup='setup_problem_1'}


```

```{r problem_1-hint}

You will need to also use the "combine" function here

```

```{r problem_1-solution}

code[c(3,7,10,13,256,1000,3000,3500)]

```

```{r problem_1-check}

grade_code()

```

### Q's 2-6: Now you are going to work with another built-in R dataset called iris. This dataset contains information about different leaf traits for different species of iris plants. In the code chunk below, use the `data()` function to load in the iris dataset. Hint: Once loaded in, R will automatically assign the name "iris" to the data frame, so use that name to reference it in subsequent questions.

```{r problem_2, exercise = TRUE}


```

```{r problem_2-hint}

Review the example for how the "faithful" data frame was loaded in

```

```{r problem_2-solution}

data(iris)

```

```{r problem_2-check}

grade_code()

```

Q3: Use the `head()` function to look at the first 3 rows of the iris data frame

```{r problem_3,exercise=TRUE,excercise.setup = "problem_2"}



```

```{r problem_3-hint}

R has automatically assigned the name "iris" to the iris data frame

```

```{r problem_3-solution}

head(x = iris,n = 3)

```

```{r problem_3-check}

grade_code()

```

Q4: Use the the `$` operator and the `mean()` function to calculate the average of the "Petal.Width" variable in the iris data frame.

```{r problem_4, exercise = TRUE, exercise.setup = "problem_2"}

```

```{r problem_4-hint}

Remember how to extract a single column?

```

```{r problem_4-solution}

mean(iris$Petal.Width)

```

```{r problem_4-check}

grade_code()

```

**Pro tip:** You can use  the `round()` function to round outputs of calculations to a specific number of decimal places.

Q5: Use the `$` operator to add a new column (variable) to the iris data frame that is calculated as the ratio of Petal.Length to Petal.Width. Name this new variable "petal_ratio".

```{r problem_5, exercise = TRUE,exercise.setup = "problem_2"}


```

```{r problem_5-hint}

Use the operator that does division!

```

```{r problem_5-solution}

iris$petal_ratio <- iris$Petal.Length/iris$Petal.Width

```

```{r problem_5-check}

grade_code()

```

### Q6: Filter the iris data frame down to the rows of data that contain petal_ratio values *greater than* the *average* of all petal_ratio values.

```{r problem_6, exercise = TRUE, exercise.setup = "problem_5-solution"}


```

```{r problem_6-hint}

Remember that when you are filtering data, you are filtering rows.

```

```{r problem_6-solution}

iris[iris$petal_ratio > mean(iris$petal_ratio),]

```

```{r problem_6-check}

grade_code()

```

### Q7: Use the same basic approach as in Q6 to find which species has the highest petal_ratio.

```{r problem_7, exercise = TRUE,exercise.setup = "problem_5-solution"}



```

```{r problem_7-hint}

Use the `max()` function

```

```{r problem_7-solution}

iris[iris$petal_ratio == max(iris$petal_ratio),]

```

```{r problem_7-check}

grade_code()

```

## Quiz 

```{r quiz_1}
quiz(
  question(
    "How would you extract the fifth value in a vector called 'secret_data'?",
    answer("secret_data(5)"),
    answer("secret_data[5]",correct = TRUE),
    answer("secret_data.5"),
    answer("secret_data{5}"))
)
```

```{r quiz_2}
quiz(
  question(
    "How would you extract the third column and fifth row from the iris data frame?",
    answer("iris[5,3]",correct = TRUE),
    answer("iris(5,3)"),
    answer("iris[3,5]"),
    answer("iris[c(3,5)]"))
  
)
```

```{r quiz_3}

#Changed to one correct answer 9/21
quiz(
  question(
    "What does the `$` operator not do?",
    answer("Add columns"),
    answer("Extract rows",correct = TRUE),
    answer("Extract columns"),
    answer("Extract an element from a list"))
  
)
```


```{r quiz_4}

#Added 9/21/26
quiz(
  question(
    "What is the correct code to return a data frame filtered all column_1 values above 4",
    answer("df[df$column_1 > 4,]",correct = TRUE)),
    answer("df[,df$column_1 > 4]"),
    answer("column_1 > 4"),
    answer("df$column_1 > 4"))
  
)
```

```{r download-answers, child = system.file("child_documents/download_answers.Rmd", package = "tutorial.helpers")}
```

All done! Please upload your html file in the correct assignment folder.