---
title: "Exam 2"
tutorial:
  id: "Week_6_Exam_2"
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
description: "Exam II focused on the basics of R"
---
<!-- File created by tutorialize.  Do not edit by hand. -->

```{r setup, include=FALSE}
library(learnr)
knitr::opts_chunk$set(echo = FALSE)
gradethis::gradethis_setup()
options(tutorial.event_recorder = tutorialize::recorder)
options(tutorial.storage = learnr::filesystem_storage(file.path(path.expand("~"), "_learnr_storage_")))
learnr::tutorial_options(exercise.startover = FALSE)
```

## Directions

This test will be 70 points and will be taken within this R tutorial. The test will be a mix of multiple choice, free response, and writing R code in response to prompts, much like you have been doing the past two weeks. You will submit the exam on D2L. You will be graded on correctness for every answer and partial credit will be given when warranted.

## Free Response (4 points total)

```{r interpretation_quiz,exercise.startover = FALSE}
quiz(
  question_text(
    "Q0: What operator is your first line of defense if you are unsure about what a function is or does or how it works in R? (1 point)",
    answer_fn(tutorialize::submitToTA),
    rows = 3,
    correct = "Your answer has been submitted"),
  question_text(
    "Q1: What are the three characteristics of tidy data? (3 points)",
    answer_fn(tutorialize::submitToTA),
    rows = 3,
    correct = "Your answer has been submitted"))  
```

## Multiple Choice (1 point each, 13 points total)

*When you submit a response, there will be a red box logging the submission and a blue box confirming it has been submitted. You can only submit once*

```{r quiz_vectors}

learnr::quiz(caption = '',
  learnr::question("Q2: What is the purpose of the `<-` operator?",
    answer("to prevent a line of code from running",correct = TRUE),
    answer("add two values",correct = TRUE),
    answer("compare two values",correct = TRUE),
    answer("assign a name to a value",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
    
  ),
   question("Q3: What is the purpose of the `!=` operator?",
    answer("indicate complete equality",correct = TRUE),
    answer("create a sequence of numbers",correct = TRUE),
    answer("Indicate inequality",correct = TRUE),
    answer("extract a value from a sequence",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
  ),
  question("Q4: True or False: all values in an atomic vector are composed of the same data type",
    answer("True",correct = TRUE),
    answer("False",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
  ),
  question("Q5: What does the c() function do?",
    answer("prevents a line of code from running",correct = TRUE),
    answer("adds two vectors together",correct = TRUE),
    answer("combines values in a sequence to create a vector",correct = TRUE),
    answer("creates a sequence of values according to a range and interval",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
  ),
   question("Q6: If you add 1:7 to c(0,1), what will the fifth value be?",
    answer("5",correct = TRUE),
    answer("6",correct = TRUE),
    answer("7",correct = TRUE),
    answer("0",correct = TRUE),
    answer("1",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
  ),
  question("Q7: In the question above, what concept is this an example of?",
    answer("Concatinating",correct = TRUE),
    answer("Recycling",correct = TRUE),
    answer("An operation",correct = TRUE),
    answer("A function",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
  ),
  question("Q8: If you want to create a vector from 1 to five million in intervals of 1000, what would be the most efficient way to do so in R?",
    answer("c() function",correct = TRUE),
    answer(": operator",correct = TRUE),
    answer("seq() function",correct = TRUE),
    answer("the function() function",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
    ),
  question(
    "Q9: Which characteristic(s) of tidy data is always an atomic vector?",
    answer("a variable/column",correct = TRUE),
    answer("a row/observation",correct = TRUE),
    answer("a cell/value",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""
    ),
  question(
    "Q10: How would you extract the 50th and 60th value in a vector called 'secret_data'?",
    answer("secret_data(50,60)",correct = TRUE),
    answer("secret_data[50,60]",correct = TRUE),
    answer("secret_data$50$60",correct = TRUE),
    answer("secret_data[c(50,60)]",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""),
  question(
    "Q11: How would you extract the first column and third row from the iris dataframe?",
    answer("iris[3,1]",correct = TRUE),
    answer("iris(1,3)",correct = TRUE),
    answer("iris[1,3]",correct = TRUE),
    answer("iris[c(3,1)]",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""),
  question(
    "Q12: What is the function(s) of the `$` operator?",
    answer("Add columns",correct = TRUE),
    answer("Extract rows",correct = TRUE),
    answer("Extract columns",correct = TRUE),
    answer("Extract from  avector",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""),
  question(
    "Q13: True or False: the `[]` operator only works on vectors",
    answer("True",correct = TRUE),
    answer("False",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = ""),
  question(
    "Q14: True or False: the `$` operator only works on columns in dataframes",
    answer("True",correct = TRUE),
    answer("False",correct = TRUE),
    allow_retry = F,
    message = "Your answer has been submitted",
    incorrect = "",
    correct = "")
)

```

## Coding (53 points total)

Q15: Create a numeric vector from 1 to 1000 that increases by values of 5. Name this vector "vector_1". (5 points)

```{r problem_1, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_1-solution}
vector_1 <- seq(from = 1,to = 1000,by = 5)
```

```{r problem_1-code-check}
grade_code()
```

Q16: Use recycling to reduce every value in vector_1 by 50%. Name this new vector "vector_1_halved". Then, create an x-y plot with vector_1 on the x axis and vector_1_halved on the y axis. (10 points)

```{r problem_2-setup}
vector_1 <- seq(from = 1, to = 1000, by = 5)
```

```{r problem_2, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_2-solution}
vector_1_halved <- vector_1*0.5
plot(vector_1,vector_1_halved)
```

```{r problem_2-code-check}
grade_code()
```


Q17: Create a dataframe that is composed of three columns: 1) a column named 'Production' that is a numeric vector composed of the values 750,500,and 250, 2) a column named 'Precipitation' that is a numeric vector composed of the values 1200, 700, and 200, and 3) a column named 'Biome' that is a character vector composed of the values "Forest","Grassland", and "Desert". Assign this new dataframe the name "NPP". (20 points)

```{r problem_3, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_3-solution}
NPP <- data.frame(
  Production = c(750,500,250),
  Precipitation = c(1200,700,200),
  Biome = c("Forest","Grassland","Desert")
)
```

```{r problem_3-code-check}
grade_code()
```

Q18: Use an extraction operator to preview the first row of the NPP dataframe you created in Q17 (while retaining all column names). (2 points)

```{r problem_4-setup}
NPP <- data.frame(Production = c(750, 500, 250), Precipitation = c(1200, 
    700, 200), Biome = c("Forest", "Grassland", "Desert"))
```

```{r problem_4, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_4-solution}
NPP[1,]
```

```{r problem_4-code-check}
grade_code()
```

Q19: Use an extraction operator to add a new column to the NPP dataframe called 'RUE' that is calculated as production divided by precipitation. This variable describes how efficiently rainfall is used for plant growth. (5 points)

```{r problem_5-setup}
NPP <- data.frame(Production = c(750, 500, 250), Precipitation = c(1200, 
    700, 200), Biome = c("Forest", "Grassland", "Desert"))
```

```{r problem_5, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_5-solution}
NPP$RUE <- NPP$Production/NPP$Precipitation
```

```{r problem_5-code-check}
grade_code()
```

Q20: Using the NPP dataframe, create an x-y plot that depicts how RUE changes as *a function of* precipitation. You can use *either* of the two approaches (but, within the same function!) to do this that were discussed in class. (5 points)

```{r problem_6-setup}
NPP <- data.frame(Production = c(750, 500, 250), Precipitation = c(1200, 
    700, 200), Biome = c("Forest", "Grassland", "Desert"))
NPP$RUE <- NPP$Production/NPP$Precipitation
```

```{r problem_6, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_6-solution}
plot(RUE ~ Precipitation,data = NPP)
#plot(NPP$Precipitation,NPP$Production) is also valid
```

```{r problem_6-code-check}
grade_code()
```

Q21: Use the data() function to load in the iris dataframe that comes with R and then use both extraction and comparison operators to filter the iris dataframe to *just* the setosa species. Name this new dataframe iris_setosa. (6 points)
```{r problem_7, exercise=TRUE, exercise.reveal_solution=FALSE}

```

```{r problem_7-solution}
data(iris)
iris_setosa <- iris[iris$Species == 'setosa',]
```

```{r problem_7-code-check}
grade_code()
```

## Submitting your Exam

Great job!  You've completed the second Exam!  Now it's time to turn it in...

In the console, enter the following (case sensitive):

```{r, echo=TRUE, eval=FALSE}
ENSC311::submit_ENSC311()
```

Your should see something like this:

```{r}
"Week_6_Exam_2_311_<Posit_Cloud_ID>"
```

where the end of the character string is an ID that Posit Cloud assigned you. As you do more tutorials, this list of tutorials you have run will get longer. You must pass the number that corresponds to this assignments. For example, *if* you wanted to turn in the first item on the list, you would enter: ENSC311::submit_ENSC311(1).

A "submit_ENSC311_..." file will appear under the "Files" tab in the lower right pane of RStudio.  Check the box next to the submit file and choose "Export..." from the "More" dropdown on the bar at the top of the tab.  Download the file and submit it to D2L. *Do not change the name of the file*

That's it!  You're all done!
